---
title: 'Chapter Title Here'
description: 'Chapter description goes here.'
---

## Example coding exercise

```yaml
type: NormalExercise
key: e8c1edbe67
lang: python
xp: 100
skills: 2
```

This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.

`@instructions`
2222222222222222222222222222222This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.This is an example exercise.

`@hint`
3333333333333333333333333333333333333333333333333333333

`@pre_exercise_code`
```{python}

```

`@sample_code`
```{python}

```

`@solution`
```{python}

```

`@sct`
```{python}

```
